---
title: Introduction to Loss Data Analytics
description: This is a template chapter.
---

## Ex 1.1

```yaml
type: NormalExercise
lang: r
xp: 100
skills: 1
key: afc746365e
```

Do some data science.

`@instructions`

Add some instructions.

`@hint`


`@pre_exercise_code`

```{r}

```

`@sample_code`

```{r}
plot(1:10)
```

`@solution`

```{r}
plot(1:10)
```

`@sct`

```{r}

```
